# Un motor de ingesta sencillo para el dataset de vuelos

### Recuerda borrar todas las líneas que dicen `raise NotImplementedError`

El trabajo final consiste en la implementación guiada de un esqueleto de motor de ingesta para ficheros JSON recibidos diariamente con información de los vuelos que han tenido lugar en el día indicado en el nombre del fichero entre dos aeropuertos de los Estados Unidos.
* La estructura de cada JSON puede consultarse abriendo con un editor de texto cualquiera de los ficheros.
* El significado de cada una de las columnas puede consultarse en el fichero config.json que se encuentra en la carpeta config del repositorio.

Vamos a probar nuestro paquete del motor de ingesta. Antes de ejecutar este notebook, es imprescindible haber completado todo el código del motor de ingesta. 
El notebook solamente valida que el código del motor de ingesta con el cual se ha generado el paquete sea correcto. Para ello, el orden en el que debemos leer, entender y completar los ficheros del repositorio es:

1. Clase `MotorIngesta` en el fichero `motor_ingesta.py`. Sólo hay que completar la función `ingesta_fichero`
2. (Opcional) Completar los tests `test_aplana` y `test_ingesta_fichero` en el fichero `test_ingesta.py`
3. Clase `FlujoDiario` en el fichero `flujo_diario.py`. Completar primero el inicializador y luego el método `procesa_diario`. Al ir completando el código de este método, veremos que a su vez necesitamos completar las dos funciones siguientes.
4. Función `aniade_hora_utc` en el fichero `agregaciones.py`
5. (Opcional) El test de dicha función se llama  `test_aniade_hora_utc` en el fichero `test_ingesta.py`
6. Función `aniade_intervalos_por_aeropuerto` en el fichero `agregaciones.py`
7. (Opcional) El test de dicha función se llama `test_aniade_intervalos_por_aeropuerto` en el fichero `test_ingesta.py`

```python
!pip install loguru==0.7.1
```

Output:
```
Requirement already satisfied: loguru==0.7.1 in ./.venv/lib/python3.12/site-packages (0.7.1)

[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m A new release of pip is available: [0m[31;49m24.3.1[0m[39;49m -> [0m[32;49m25.2[0m
[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m To update, run: [0m[32;49mpip install --upgrade pip[0m
```

## Instalamos el wheel en el cluster. 

```python
!pip install --force-reinstall ../motor_ingesta-0.1.0-py3-none-any.whl
```

Output:
```
Processing /Users/orr21/Documents/UCM/Master/Spark/spark-tarea-final/motor_ingesta-0.1.0-py3-none-any.whl
Installing collected packages: motor-ingesta
  Attempting uninstall: motor-ingesta
    Found existing installation: motor-ingesta 0.1.0
    Uninstalling motor-ingesta-0.1.0:
      Successfully uninstalled motor-ingesta-0.1.0
Successfully installed motor-ingesta-0.1.0

[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m A new release of pip is available: [0m[31;49m24.3.1[0m[39;49m -> [0m[32;49m25.2[0m
[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m To update, run: [0m[32;49mpip install --upgrade pip[0m
```

### Probamos la ingesta de un fichero

**Ejercicio 1 (2 puntos)**. Ingestar el fichero **`2023-01-01.json`** utilizando el motor de ingesta completo. Debe crearse el objeto de la clase MotorIngesta y utilizar el método `ingesta_fichero`, dejando el resultado en la variable `flights_df`. La variable `flujo_diario` contiene un objeto FlujoDiario inicializado con el path de configuración anterior, y lo vamos a usar ahora solo para leer adecuadamente la configuración y pasársela al objeto `motor_ingesta` como el argumento config.

* Este ejercicio requiere haber completado previamente el código del paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

<!-- grade_id: ingesta_fichero, points: 0, grade: false, solution: true -->
```python
from motor_ingesta.motor_ingesta import MotorIngesta
import json

path_config_flujo_diario = "../config/config.json"       # ruta del fichero config.json, que no pertenece al paquete
path_json_primer_dia = "../data/landing/2023-01-01.json"          # ruta del fichero JSON de un día concreto que queremos ingestar, en nuestro caso 2023-01-01.json

with open(path_config_flujo_diario) as f:
    flujo_diario = json.load(f)
motor_ingesta = MotorIngesta(flujo_diario)
flights_df = motor_ingesta.ingesta_fichero(path_json_primer_dia)

flights_df.show(5)
```

Output:
```
25/09/11 16:27:03 WARN Utils: Your hostname, MacBook-Pro-de-Oscar.local resolves to a loopback address: 127.0.0.1; using 192.168.1.112 instead (on interface en0)
25/09/11 16:27:03 WARN Utils: Set SPARK_LOCAL_IP if you need to bind to another address
Setting default log level to "WARN".
To adjust logging level use sc.setLogLevel(newLevel). For SparkR, use setLogLevel(newLevel).
25/09/11 16:27:03 WARN NativeCodeLoader: Unable to load native-hadoop library for your platform... using builtin-java classes where applicable
25/09/11 16:27:05 WARN SparkStringUtils: Truncated the string representation of a plan since it was too large. This behavior can be adjusted by setting 'spark.sql.debug.maxToStringFields'.
+----------+-----------------+---------------+------+---------------+-----------+-------------+----+---------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+
|FlightDate|Reporting_Airline|OriginAirportID|Origin| OriginCityName|OriginState|DestAirpor
```

<!-- grade_id: ingesta-tests, points: 1, grade: true, solution: false -->
```python
assert(flights_df.count() == 15856)
assert(len(flights_df.columns) == 18)
dtypes = dict(flights_df.dtypes)
assert(dtypes["Diverted"] == "boolean")
assert(dtypes["ArrTime"] == "int")
assert(flights_df.schema["Dest"].metadata == {"comment": "Destination Airport IATA code (3 letters)"})
```

### Probamos la función de añadir la hora en formato UTC

**Ejercicio 2 (2 puntos)** Probar la función de añadir hora UTC con el DF `flights_df` construido anteriormente. El resultado debe dejarse en la variable `flights_with_utc`. Recuerda que esto no es propiamente un test unitario.

* Este ejercicio requiere haber completado previamente el código de la función `aniade_hora_utc` del paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

<!-- grade_id: flights-utc, points: 0, grade: false, solution: true -->
```python
from motor_ingesta.agregaciones import aniade_hora_utc
from pyspark.sql import SparkSession

spark = SparkSession.builder \
    .appName("Actividad Spark") \
    .getOrCreate()

flights_with_utc = aniade_hora_utc(spark, flights_df)

flights_with_utc.show(5)
```

Output:
```
25/09/11 16:27:13 WARN SparkSession: Using an existing Spark session; only runtime SQL configurations will take effect.
+----------+-----------------+---------------+------+---------------+-----------+-------------+----+---------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+-------------------+
|FlightDate|Reporting_Airline|OriginAirportID|Origin| OriginCityName|OriginState|DestAirportID|Dest|   DestCityName|DestState|DepDelay|DepTime|ArrDelay|ArrTime|Cancelled|Diverted|AirTime|Distance|         FlightTime|
+----------+-----------------+---------------+------+---------------+-----------+-------------+----+---------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+-------------------+
|2023-01-01|               9E|          12884|   LAN|    Lansing, MI|         MI|        11433| DTW|    Detroit, MI|       MI|      -6|   1136|     -22|   1215|    false|   false|     21|      74|2023-01-01 16:36:00|
|2023-01-01|
```

<!-- grade_id: flights-utc-tests, points: 2, grade: true, solution: false -->
```python
from pyspark.sql import functions as F
assert(flights_with_utc.where("FlightTime is null").count() == 266)
types = dict(flights_with_utc.dtypes)
assert(flights_with_utc.dtypes[18] == ("FlightTime", "timestamp"))

first_row = flights_with_utc.where("OriginAirportID = 12884").select(F.min("FlightTime").cast("string").alias("FlightTime")).first()
assert(first_row.FlightTime == "2023-01-01 10:59:00")
```

## Probamos la función de añadir las columnas con la hora del siguiente vuelo, su aerolínea y el intervalo de tiempo transcurrido

**Ejercicio 3 (2.5 puntos)** Invocar a la función de añadir intervalos por aeropuerto, partiendo de la variable `flights_with_utc` del apartado anterior, dejando el resultado devuelto por la función en la variable `df_with_next_flight` cacheada.

* Este ejercicio requiere haber completado previamente el código de la función `aniade_intervalos_por_aeropuerto` en el paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

<!-- grade_id: flights-intervalos, points: 0, grade: false, solution: true -->
```python
from motor_ingesta.agregaciones import aniade_intervalos_por_aeropuerto

df_with_next_flight = aniade_intervalos_por_aeropuerto(flights_with_utc)

df_with_next_flight.show(10, False)
```

Output:
```
+----------+-----------------+---------------+------+------------------------------+-----------+-------------+----+---------------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+-------------------+-------------------+------------+---------+
|FlightDate|Reporting_Airline|OriginAirportID|Origin|OriginCityName                |OriginState|DestAirportID|Dest|DestCityName         |DestState|DepDelay|DepTime|ArrDelay|ArrTime|Cancelled|Diverted|AirTime|Distance|FlightTime         |FlightTime_next    |Airline_next|diff_next|
+----------+-----------------+---------------+------+------------------------------+-----------+-------------+----+---------------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+-------------------+-------------------+------------+---------+
|2023-01-01|9E               |10135          |ABE   |Allentown/Bethlehem/Easton, PA|PA         |10397        |ATL |Atlanta, GA          |GA       |1       |6
```

<!-- grade_id: flights-intervalos-test, points: 2.5, grade: true, solution: false -->
```python
assert(df_with_next_flight.dtypes[19] == ("FlightTime_next", "timestamp"))
assert(df_with_next_flight.dtypes[20] == ("Airline_next", "string"))
assert(df_with_next_flight.dtypes[21] == ("diff_next", "bigint"))

first_row = df_with_next_flight.where("OriginAirportID = 12884")\
                               .select(F.col("FlightTime").cast("string"), 
                                       F.col("FlightTime_next").cast("string"), 
                                       F.col("Airline_next"),
                                       F.col("diff_next")).sort("FlightTime").first()

assert(first_row.FlightTime_next == "2023-01-01 16:36:00")
assert(first_row.Airline_next == "9E")
assert(first_row.diff_next == 20220)
```

## Corregimos que el último vuelo de cada aeropuerto y cada día tiene valor nulo en las 3 columnas `_next`

**Ejercicio 4 (2.5 puntos)**

Tal como está implementada la lógica del flujo diario, el último vuelo de cada día no tendrá informada la columna FlightTime_next porque no se dispone todavía de datos del día siguiente. Se pide **corregir este comportamiento** para solucionar los valores nulos, modificando el código del método `procesa_diario` de manera que, antes de escribir los datos del día actual, se hayan corregido las tres columnas `_next` en los datos del día anterior al que estamos ingestando. Una manera simple (aunque no necesariamente óptima) de conseguirlo es:
* Leer de la tabla la partición que se escribió el día previo, si existiera dicha tabla y dicha partición.
* Añadir al DF devuelto por `aniade_hora_utc` las 3 columnas que le faltan para tener la misma estructura que la tabla, que son `FlightTime_next`, `Airline_next` y `diff_next` (pueden ser en ese orden si la función `aniade_intervalos_por_aeropuerto` se ha implementado para añadirlas en ese orden), pero sin darles valor (con valor None, convirtiendo cada columna al tipo de dato adecuado para que después encaje con la tabla existente).
* Unir el DF del día previo y el que acabamos de calcular
* Invocar a `aniade_intervalos_por_aeropuerto` pasando como argumento el DF resultante de la unión.

Aparte de un test unitario (que se deja como optativo pero sin puntuación), la manera de comprobar el funcionamiento será invocar a `procesa_diario` del flujo diario, con los ficheros de dos días consecutivos, y después comprobar lo que se ha escrito en la tabla tras la ingesta del segundo fichero. Lo probaremos con los días 1 y 2 de enero de 2023.

* Este ejercicio requiere haber completado previamente el código del paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

```python
from motor_ingesta.flujo_diario import FlujoDiario

ruta_config = "../config/config.json"
path_json_primer_dia = "../data/landing/2023-01-01.json"

flujo = FlujoDiario(ruta_config)
flujo.procesa_diario(path_json_primer_dia)
```

Output:
```
[32m2025-09-11 16:27:16.753[0m | [1mINFO    [0m | [36mmotor_ingesta.flujo_diario[0m:[36mprocesa_diario[0m:[36m46[0m - [1mNo se han podido leer datos del día 2022-12-31: [TABLE_OR_VIEW_NOT_FOUND] The table or view `default`.`flights` cannot be found. Verify the spelling and correctness of the schema and catalog.
If you did not qualify the name with a schema, verify the current_schema() output, or qualify the name with the correct schema and catalog.
To tolerate the error on drop use DROP VIEW IF EXISTS or DROP TABLE IF EXISTS.;
'UnresolvedRelation [default, flights], [], false
[0m
25/09/11 16:27:16 WARN GarbageCollectionMetrics: To enable non-built-in garbage collector(s) List(G1 Concurrent GC), users should configure it(them) to spark.eventLog.gcMetrics.youngGenerationGarbageCollectors or spark.eventLog.gcMetrics.oldGenerationGarbageCollectors
```

<!-- grade_id: procesa-diario, points: 0, grade: false, solution: true -->
```python
path_json_segundo_dia = "../data/landing/2023-01-02.json"  # path del fichero 2023-01-02.json

flujo.procesa_diario(path_json_segundo_dia)
```

Output:
```
[32m2025-09-11 16:27:18.565[0m | [1mINFO    [0m | [36mmotor_ingesta.flujo_diario[0m:[36mprocesa_diario[0m:[36m44[0m - [1mLeída partición del día 2023-01-01 con éxito[0m
```

<!-- grade_id: tests-unitarios, points: 2.5, grade: true, solution: false -->
```python
vuelos = spark.read.table("default.flights").sort("Origin", "FlightTime")
assert(vuelos.count() == 33931)
row = vuelos.where("FlightDate = '2023-01-01' and  Origin = 'ABE' and DepTime = 1734").first()
assert(row.diff_next == 44220)
```

## Ejercicio opcional

**Ejercicio opcional (1 punto)** Completa los cuatro tests unitarios que encontrarás en el fichero `test_ingesta.py`. No tienes que escribir más código en este notebook.

- Se podrá optar a una calificación final de hasta 9.0 puntos sin resolver este ejercicio.